# Clustering Algorithms (20 points) [TA: Vodelina Samatova]

In this exercise, you'll be working on different clustering algorithms: KMeans and Hierarchical Clustering Algorithms.   


```python
import numpy as np
np.random.seed(23)
from sklearn.metrics import pairwise_distances
from sklearn.cluster import KMeans, AgglomerativeClustering
import matplotlib.pyplot as plt

data_folder = 'data'
```

```python
# First, let us load datasets to use for KMeans and Hierarchical Clustering
d1 = np.load(f'{data_folder}/hw4-p1-d1.npz')
dataset_1 = d1['data']
print(f'Dataset 1 is of dimension {dataset_1.shape}')
initial_centers = d1['initial']

d2 = np.load(f'{data_folder}/hw4-p1-d2.npz')
dataset_2 = d2['data']
print(f'Dataset 2 is of dimension {dataset_2.shape}')
initial_centers = d2['initial']

d3 = np.load(f'{data_folder}/hw4-p2-d2.npz')
dataset_3 = d3['data']
print(f'Dataset 3 is of dimension {dataset_2.shape}')

# plot the data 
plt.scatter(x = dataset_1[:, 0], y=dataset_1[:, 1])
plt.title('Dataset 1')
plt.show()

plt.scatter(x = dataset_2[:, 0], y=dataset_2[:, 1])
plt.title('Dataset 2')
plt.show()

plt.scatter(x = dataset_3[:, 0], y=dataset_3[:, 1])
plt.title('Dataset 3')
plt.show()
```

# Problem 1.1: Cluster SSE

<!-- grade_id: alda_calculate_sse, points: 0, grade: false, solution: true -->
```python
def alda_calculate_sse(data, cluster_assignments):
    """
    TODO: Write code to calculate the SSE for the dataset 'data' given cluster assignments. Code should be written from
    scratch, not use an existing sklearn library to do so (numpy is allowed to use).
    
    INPUT:
          data: the dataset of size (n_samples, n_attributes) where n_samples is the number of samples
                and n_attributes is the number of attributes for each data point.
          cluster_assignments: a numpy integer vector containing the cluster assignments for each point in data.
                           It is of shape (n_attributes, ).
          
    OUTPUT:
          A floating point value containing the SSE
    
    Allowed functions: numpy mathematical functions
    """
     
```

```python
# Test your code!
# assume we're creating a dataset with two clusters 
practice_data = np.array([[0, 0], [1, 1], [2, 2], [3, 3] ,[4, 4]])
# here, 0 means that all points belong to cluster 0
# 1 means all points belong to cluster 1
practice_cluster_assignments = np.array([0, 0, 1, 1, 1]) 

# Let's calculate its SSE
print(f'For given data, SSE = {alda_calculate_sse(practice_data, practice_cluster_assignments)}')
# should be 5

plt.scatter(x = practice_data[:, 0], y=practice_data[:, 1], c=practice_cluster_assignments)
plt.title('Dataset 2')
plt.show()
```

<!-- grade_id: alda_calculate_sse-public, points: 6, grade: true, solution: false -->
```python
# Let us evaluate the function on the practice dataset 
practice_ds = np.load(f'{data_folder}/hw4-p1-practice.npz')
practice_data = practice_ds['data']
practice_initial_centers = practice_ds['initial']
practice_sse = alda_calculate_sse(practice_data, practice_initial_centers)
np.testing.assert_almost_equal(practice_sse, 2554.26922225574)
```

<!-- grade_id: alda_calculate_sse-private, points: 4, grade: true, solution: false -->
```python
# Remember, we will have hidden tests as well!
```

# Problem 1.2: K-Means Clustering

In this problem, you will be using the [KMeans](https://scikit-learn.org/stable/modules/generated/sklearn.cluster.KMeans.html) sklearn **library**,  to perform k-means clustering.

<!-- grade_id: alda_kmeans, points: 0, grade: false, solution: true -->
```python
def alda_kmeans(data, k, initial_centers, max_iter):
    """
    TODO: Write code to perform k-means clustering on the data matrix given by the 'data' variable using the kmeans. 
    
    INPUT:
        data: A 2D numpy array of shape (n_samples, n_attributes) where n_samples is the number of samples
              and n_attributes is the number of attributes for each data point.
        k: An integer, gives the number of clusters to calculate in kmeans
        initial_centers: A numpy array of shape (k, n_attributes) containing the initial centers
        max_iter: max. number of iterations to run for k-means
    
    OUTPUT:
        A numpy ndarray of shape (n_samples, ) providing the cluster assignments
        Cluster assignment values range between 0 and k-1.
    
    NOTE 1: Scikit-learn KMeans Reference: https://scikit-learn.org/stable/modules/generated/sklearn.cluster.KMeans.html
    NOTE 2: For reproducibility, use the initial_centers variable for the init agrument in KMeans, 
            specifying initial centers; you'll also have to set n_init = 1 to remove any warnings.
    """

```

```python
# Test your code!
# let us initialize some variables first
practice_k = 2 # to represent number of clusters

practice_ds = np.load(f'{data_folder}/hw4-p1-practice.npz')
# the dataset is stored here
practice_data = practice_ds['data']
# initial cluster assignments are stored in this variable
# We use these to ensure a deterministic answer (for testing purposes)
practice_initial_assignments = practice_ds['initial']
# calculate initial centers
practice_centers = []
for practice_c in np.unique(practice_initial_assignments):
    p_current_cluster_data = practice_data[practice_initial_assignments == practice_c, : ]
    p_current_cluster_centers = np.mean(p_current_cluster_data, axis=0)
    practice_centers.append(p_current_cluster_centers)

# convert practice_centers to a numpy array 
practice_centers = np.array(practice_centers)

# call kmeans
practice_assignments = alda_kmeans(data=practice_data, k=practice_k,\
                                       initial_centers=practice_centers, max_iter=2)

print(f'SSE of KMeans on practice dataset is \
{alda_calculate_sse(practice_data, practice_assignments)}') # Should be ~ 133.045
```

<!-- grade_id: alda_kmeans-public, points: 2, grade: true, solution: false -->
```python
# Public Test Cases
# let us initialize some variables first
test_k = 3 # to represent number of clusters

test_ds = np.load(f'{data_folder}/hw4-p1-d1.npz')
# the dataset is stored here
test_data = test_ds['data']
# initial cluster assignments are stored in this variable
test_initial_assignments = test_ds['initial']

# calculate initial centers
test_centers = []
for test_c in np.unique(test_initial_assignments):
    test_cluster_data = test_data[test_initial_assignments == test_c, : ]
    test_cluster_centers = np.mean(test_cluster_data, axis=0)
    test_centers.append(test_cluster_centers)

# convert practice_centers to a numpy array 
test_centers = np.array(test_centers)

# call kmeans
test_assignments = alda_kmeans(data=test_data, k=test_k,\
                                       initial_centers=test_centers, max_iter=4)


np.testing.assert_almost_equal(alda_calculate_sse(test_data, test_assignments), 529.2865169916909)
```

<!-- grade_id: alda_kmeans-prviate, points: 2, grade: true, solution: false -->
```python
# Remember, we will have hidden test cases too
```

Below, we have some code to visualize your k-means clustering on the three toy datasets:

```python
def kmeans(k,ds, title='Plot'):
    # the dataset is stored here
    data = ds['data']
    # initial cluster assignments are stored in this variable
    initial_assignments = ds['initial']

    # calculate initial centers
    centers = []
    for c in np.unique(initial_assignments):
        cluster_data = data[initial_assignments == c, : ]
        cluster_centers = np.mean(cluster_data, axis=0)
        centers.append(cluster_centers)

    # convert practice_centers to a numpy array 
    centers = np.array(centers)
    assignment = alda_kmeans(data=data, k=k, initial_centers=centers, max_iter=300)

    plt.scatter(x = data[:, 0], y=data[:, 1], c=assignment)
    plt.title(title)
    plt.show()
```

```python
ds1 = np.load(f'{data_folder}/hw4-p1-d1.npz')
assignment1 = kmeans(3,ds1, 'Blobs')
ds2 = np.load(f'{data_folder}/hw4-p1-d2.npz')
assignment2 = kmeans(2,ds2, 'Two moons')
ds3 = np.load(f'{data_folder}/hw4-p2-d2.npz')
assignment3 = kmeans(2,ds3, 'Two moons (loose)')
```

## 1.2.1: k-means++ (Example)

[k-means++](https://en.wikipedia.org/wiki/K-means%2B%2B) is a clustering algorithm with the optimized centroids initialization, fortunately, the built-in [KMeans](https://scikit-learn.org/stable/modules/generated/sklearn.cluster.KMeans.html) function in sklearn applied k++ centroid initialization as a default, we will compare how many iterations are required to get convergence with and without using k++ algorithm.

```python
def kmeans_kplusplus(k,ds):
    # the dataset is stored here
    data = ds['data']
    kmeans = KMeans(n_clusters=k, n_init=1, max_iter=300, verbose=1)
    return kmeans.fit(data).labels_
```

```python
print('*'*20 + "  See how many iterations are required without k++ initialization  " + '*'*20)
assignment_without_kplusplus = kmeans(3,ds1)
print('*'*20 + "  See how many iterations are required with k++ initialization  " + '*'*20)
assignment_with_kplusplus = kmeans_kplusplus(3,ds1)
```

# 1.3. Agglomerative/Hierarchical Clustering

You need to implement the same function, agglomerative clusting using the [AgglomerativeClustering](https://scikit-learn.org/stable/modules/generated/sklearn.cluster.AgglomerativeClustering.html) library from sklearn.

<!-- grade_id: alda_agglomerative, points: 0, grade: false, solution: true -->
```python
def alda_agglomerative(data, k, linkage):
    """
    TODO: Perform agglomerative clustering on the data from the 'data' variable. Type of linkage is 
    specified by 'linkage'
    INPUT:
         data: A 2D numpy array of shape (n_samples, n_attributes) where n_samples is the number of samples
              and n_attributes is the number of attributes for each data point.
         k: An integer, gives the number of clusters to calculate 
         linkage: A string, can be one of the three: 'single', 'complete', 'average'
         
    OUTPUT:
         A numpy ndarray of shape (n_samples, ) providing the cluster assignments. 
         Cluster assignment values range between 0 and k-1.
    
         
    """     
```

```python
# Test your code!
# let us initialize some variables first
practice_k = 2 # to represent number of clusters

practice_ds = np.load(f'{data_folder}/hw4-p1-d2.npz')
# the dataset is stored here
practice_data = practice_ds['data']

# perform hierarchical clustering using complete linkage
linkage_type = 'complete'
practice_cluster_assignments = alda_agglomerative(practice_data, practice_k, linkage_type)

print(f'SSE of agglomerative clustering with {linkage_type} linkage \
is {alda_calculate_sse(practice_data, practice_cluster_assignments)}') # Should be ~ 81.539

# plot the cluster assignments
plt.scatter(practice_data[:, 0], practice_data[:, 1], c=practice_cluster_assignments)

```

<!-- grade_id: alda_agglomerative-public, points: 2, grade: true, solution: false -->
```python
# Public Tests
test_k = 2 # to represent number of clusters

test_ds = np.load(f'{data_folder}/hw4-p1-d2.npz')
# the dataset is stored here
test_data = test_ds['data']

# perform hierarchical clustering using single linkage
linkage_type = 'single'
test_cluster_assignments = alda_agglomerative(test_data, test_k, linkage_type)

np.testing.assert_almost_equal(alda_calculate_sse(test_data, test_cluster_assignments), 121.76952070250866)

# plot the clusters
plt.scatter(test_data[:, 0], test_data[:, 1], c=test_cluster_assignments)

```

<!-- grade_id: alda_agglomerative-private, points: 2, grade: true, solution: false -->
```python
# Remember, we'll have hidden test cases too
```

### 1.3.1: Dendrogram

SKlearn doesn't have a built-in method for dendograms, so we'll be using [scipy's dendrogram tool](https://docs.scipy.org/doc/scipy/reference/generated/scipy.cluster.hierarchy.dendrogram.html)

Let's try the dendrogram on the `dataset_2` we had above.

```python
import scipy.cluster.hierarchy as hierarchy

# Note the method='single'
single_link = hierarchy.linkage(dataset_2, method='single')
hierarchy.dendrogram(single_link)
plt.show()
```

Now, using the same logic as the above code, implement the complete-link dendrogram.

<!-- grade_id: complete_link, points: 0, grade: false, solution: true -->
```python
# Note that it is REQUIRED to use complete_link_clustering and dendrogram 
# variables names when creating your clustering and dendrogram

#TODO: implement the complete-link dendrogram

complete_link_clustering = None
dendrogram = None

```

<!-- grade_id: complete_link-public, points: 2, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(dendrogram['icoord'][5][2],34.375)
np.testing.assert_almost_equal(complete_link_clustering[24][2],0.017184987030773195)
```

### Interpreting the dendrograms

By looking at the dendrograms, what **difference** do you see between single-link clustering and complete-link? Which one has performed better?

### Answer here


# 1.4. DBSCAN
Now we'll look at a dataset where neither K-means nor Agglomerative Clustering does well.

### Example 4.1: Limitations of k-means and hierarchical clustering
First, let's take another look at the scatter plot of `dataset_3`.

```python
plt.scatter(x = dataset_3[:, 0], y=dataset_3[:, 1])
plt.title('Dataset 3')
plt.show()
```

#### Trying K-means
Let's try a K-means on this dataset with `n_clusters=2` and the rest of the same paramaters used previously.

```python
# Try K-means
kmeans = KMeans(n_clusters=2, init='random', n_init=1, max_iter=10, verbose=1)
results = kmeans.fit(dataset_3)
plt.scatter(x = dataset_3[:, 0], y=dataset_3[:, 1], c=results.labels_)
```

#### Trying Agglomerative
Now let's try a agglomerative clustering on this dataset with `n_clusters=2`. Try it with both `complete` and `single` linkage.

```python
# Try agglomerative clustering
agglom = AgglomerativeClustering(n_clusters=2, linkage="single")
results = agglom.fit(dataset_3)
plt.scatter(x = dataset_3[:, 0], y=dataset_3[:, 1], c=results.labels_)
```

```python
# Try agglomerative clustering
agglom = AgglomerativeClustering(n_clusters=2, linkage="complete")
results = agglom.fit(dataset_3)
plt.scatter(x = dataset_3[:, 0], y=dataset_3[:, 1], c=results.labels_)
```

Now is the time to try DBScan on `dataset_3`. Will it make better clusters? We don't know! Let's see

### Problem 1.4.1: DBSCAN

Finding the right number for `eps` and `min_samples` (min points) can be a bit touchy. **A good rule of thumb is to choose `min_samples` to be TWICE the dimension of the feature space.**

As the first step, complete the function `plot_dbscan`. The goal is to have a function which gets `epsilon` and `min_samples` and performs DBScan clustering on `dataset_3`. Using [the sklearn documentation for DBSCAN](https://scikit-learn.org/stable/modules/generated/sklearn.cluster.DBSCAN.html) can be so helpful.

<!-- grade_id: dbscan, points: 0, grade: false, solution: true -->
```python
from sklearn.cluster import DBSCAN

#TODO: complete the following function
def plot_dbscan(epsilon, min_samples):
    # Output:
    #   clustering: a fitted instance of DBScan clustering
    #   plt.scatter: scatter plot of DBScan
    x = dataset_3[:, 0]
    y = dataset_3[:, 1]
    clustering = None

    return clustering, plt.scatter(x=x, y=y, c=clustering.labels_)
```

<!-- grade_id: dbscan-public, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(plot_dbscan(0.1, 4)[0].labels_[5],-1)
```

Now let's play around with different epsilons and see how it affects the result of DBScan clustering. What version of eps worked best for you?

```python
# Since we're in two-dimensional space, we will set min points as 2*2=4
min_pts = 4

#Try with different eps from 0.05 to 0.25
plot_dbscan(0.5, min_pts)
```

What value of eps worked best? How does eps affect your clustering? ANSWER HERE

### Problem 1.4.2: Finding a Reasonable Value for Epsilon in DBSCAN
In the real world, we cannot find a good epsilon based on guessing. So we'll be using the method discussed in class to determine a good `eps` for a given `min_samples`. That is, if `min_samples=k`, plot the distance from each point to its kth nearest neighbor. Then, sort those distances and plot them. Finally, locate on the graph where the sharpest change begins, and use that distance as your `eps`. You might have to play with values near that to get a good clustering.

```python
from sklearn.neighbors import NearestNeighbors

#Get the nearest neighbors
# the reason why we set min_pts as 4 is discussed above
min_pts = 4
neighbors = NearestNeighbors(n_neighbors=min_pts).fit(dataset_3)
distances, indices = neighbors.kneighbors(dataset_3)
#Sort them
dis_to_nth_nn = sorted(distances[:,min_pts-1])
#Plot
plt.plot(dis_to_nth_nn)
plt.title(f'(Identifying best Eps value for MinPts={min_pts}')
plt.xlabel("nth Greatest Datapoint")
plt.ylabel(f'Distance to {min_pts}th Nearest Neighbor')
```

As we can see, the sharpest cange begins at around the 175th greatest datapoint, which corresponds to a value of around `0.15`. So that can serve as a starting point for trying out different `eps` values.

```python
# Now try different eps from 0.15-0.22. 0.21 seems to be the best

# plot_dbscan(0.15, min_pts)
```

As you can see, for some datasets, the output of DBScan can be way more acceptable than k-means and hierarchical clustering. We just have to find good values of epsilon and min_samples.

```python

```